[1643806465] SERVICE ALERT: localhost;Total Processes;CRITICAL;SOFT;1;CRITICAL - Le plugin n'as pas répondu dans les 10 secondes
[1643806513] SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 106 processus avec ETAT = RSZDT
[1643806672] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - Charge moyenne: 6.18, 14.01, 7.72
[1643806732] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;2;CRITICAL - Charge moyenne: 2.78, 11.62, 7.29
[1643806792] SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;3;CRITICAL - Charge moyenne: 1.30, 9.61, 6.87
[1643806852] SERVICE ALERT: localhost;Current Load;CRITICAL;HARD;4;CRITICAL - Charge moyenne: 1.34, 8.16, 6.55
[1643807452] SERVICE ALERT: localhost;Current Load;WARNING;HARD;4;WARNING - Charge moyenne: 0.77, 1.98, 3.95
[1643807752] SERVICE ALERT: localhost;Current Load;OK;HARD;4;OK - Charge moyenne: 0.36, 0.92, 2.94
[1643839852] SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - Charge moyenne: 7.29, 4.33, 2.27
[1643839912] SERVICE ALERT: localhost;Current Load;OK;SOFT;2;OK - Charge moyenne: 2.75, 3.56, 2.13
[1643842110] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% libre (0 MB sur un total de 3711 MB)
[1643842170] SERVICE ALERT: localhost;Swap Usage;OK;SOFT;2;SWAP OK - 1% libre (0 MB sur un total de 3711 MB)
[1643842470] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% libre (0 MB sur un total de 3711 MB)
[1643842530] SERVICE ALERT: localhost;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% libre (0 MB sur un total de 3711 MB)
[1643842590] SERVICE ALERT: localhost;Swap Usage;OK;SOFT;3;SWAP OK - 1% libre (0 MB sur un total de 3711 MB)